rule_files:
  - underwriting-promtool.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'api_prediction_entropy'
        values: '0.8 0.8'
    alert_rule_test:
      - eval_time: 1m
        alertname: PredictionEntropyTooHigh
        exp_alerts:
          - exp_labels:
              alertname: PredictionEntropyTooHigh
              severity: warning
            exp_annotations:
              summary: "⚠️ Prediction entropy too high: 0.8"

  - interval: 1m
    input_series:
      - series: 'api_avg_confidence'
        values: '0.5 0.5'
    alert_rule_test:
      - eval_time: 1m
        alertname: ConfidenceTooLow
        exp_alerts:
          - exp_labels:
              alertname: ConfidenceTooLow
              severity: warning
            exp_annotations:
              summary: "⚠️ Model confidence low: 0.5"
